# 🧱 Python Practice: Data Structures — Lists and Dictionaries

Welcome to this practice notebook on **Python Data Structures**, with a strong focus on **lists** and **dictionaries** — two foundational tools that power everything from simple scripts to complex data processing workflows.

This notebook is designed to strengthen your understanding of how to create, manipulate, and extract value from these data structures. You'll go beyond basic operations and work on real-world patterns such as frequency counting, nested lookups, inventory updates, leaderboard creation, and list transformation.

You’ll also explore useful tools from the `collections` module like `defaultdict`, `Counter`, and `deque`, which make certain tasks much easier and more efficient.

---

## 🧭 Instructions
- This notebook contains **10 mandatory exercises** and **3 optional (bonus) exercises** that are a bit more challenging.
- Each task focuses on a key list or dictionary concept, such as merging, flattening, grouping, or nested updates.
- Hints are available if you get stuck — take your time and explore!
- You can run the tests as many times as needed.
- Total estimated time: **60 minutes**.
- Think like a problem solver: model, organize, and transform data clearly and efficiently.

---

## 🧠 Concepts Covered
- List slicing, filtering, reversing, and rotating
- Dictionary creation, lookup, merging, and conflict resolution
- Working with nested data structures
- Grouping and frequency counting using `defaultdict` and `Counter`
- Deque-based rotation and efficient list handling
- Real-world scenarios like building leaderboards and updating inventories
- Writing clean, reusable Python functions using lists and dictionaries

---

Let’s dive in!

## Exercise 1: Reverse a List

**Task**: Write a function `reverse_list(lst)` that takes a list and returns a **new list** with the elements in reverse order.

📌 Do **not** use `lst[::-1]`. Use a loop or built-in methods.

<details>
  <summary><strong>Click here for a hint!</strong></summary>
  Use a loop or try <code>list(reversed(lst))</code>.
</details>

<!-- grade_id: cell-317cc3d571cc1a2b, points: 0, grade: false, solution: true -->
```python
def reverse_list(lst):
    new_list = []
    for item in lst:
      new_list.insert(0, item)
    return new_list

```

<!-- grade_id: cell-b1811b43f1452005, points: 1, grade: true, solution: false -->
```python
# Tests
assert reverse_list([1, 2, 3]) == [3, 2, 1]
assert reverse_list(["a", "b", "c"]) == ["c", "b", "a"]
```

## Exercise 2: Remove Duplicates from a List

**Task**: Write a function `remove_duplicates(lst)` that returns a **new list** with all duplicates removed, keeping only the **first occurrence** of each element.

📌 The original order must be preserved.

<details>
  <summary><strong>Click here for a hint!</strong></summary>
  Use a loop and a helper set to track seen items.
</details>

<!-- grade_id: cell-e0acce872ea69156, points: 0, grade: false, solution: true -->
```python
def remove_duplicates(lst):
  unique_list = []
  for item in lst:
    if item not in unique_list:
      unique_list.append(item)
  return unique_list

```

<!-- grade_id: cell-78f1cf0f3f5a3547, points: 1, grade: true, solution: false -->
```python
# Tests
assert remove_duplicates([1, 2, 2, 3, 1]) == [1, 2, 3]
assert remove_duplicates(["a", "a", "b"]) == ["a", "b"]
```

## Exercise 3: Nested Dictionary Lookup

**Task**: Write a function `nested_lookup(d, outer_key, inner_key)` that looks up a value in a **nested dictionary** like this:

```python
{
  "user1": {"age": 15, "grade": "10th"},
  "user2": {"age": 16, "grade": "11th"}
}

<details>
  <summary><strong>Click here for a hint!</strong></summary>
  Use <code>.get()</code> twice: once on the outer dictionary, then on the inner one.
</details>

<!-- grade_id: cell-1a0bec2c7c3d8093, points: 0, grade: false, solution: true -->
```python
def nested_lookup(d, outer_key, inner_key):
    result = d.get(outer_key, {}).get(inner_key, "Not found")
    return result
```

<!-- grade_id: cell-39750eb8433c0e1f, points: 1, grade: true, solution: false -->
```python
# Tests
users = {
    "alice": {"age": 15, "grade": "10th"},
    "bob": {"age": 16, "grade": "11th"}
}
assert nested_lookup(users, "alice", "age") == 15
assert nested_lookup(users, "bob", "grade") == "11th"
assert nested_lookup(users, "charlie", "age") == "Not found"
assert nested_lookup(users, "bob", "email") == "Not found"
```

## Exercise 4: Sum of Even Numbers in a List

**Task**: Write a function `sum_even(lst)` that returns the **sum of all even numbers** in the list `lst`.

📌 If there are no even numbers, return 0.

<details>
  <summary><strong>Click here for a hint!</strong></summary>
  Use a loop or list comprehension to filter even numbers, then use <code>sum()</code>.
</details>

<!-- grade_id: cell-1d98f487541e0fdf, points: 0, grade: false, solution: true -->
```python
def sum_even(lst):
    total = 0
    for num in lst:
        if num % 2 == 0:
            total += num
    return total
```

<!-- grade_id: cell-6cdb04262631ad15, points: 1, grade: true, solution: false -->
```python
# Tests
assert sum_even([1, 2, 3, 4, 5]) == 6
assert sum_even([1, 3, 5]) == 0
```

## Exercise 5: Merge Dictionaries with Conflict Resolution

**Task**: Write a function `smart_merge(d1, d2)` that merges two dictionaries, but:

- If a key is in **both** dictionaries, store the values in a **list**.
- Otherwise, store the single value.

📌 Example:
```python
smart_merge({"a": 1}, {"a": 2}) ➞ {"a": [1, 2]}
smart_merge({"x": 5}, {"y": 6}) ➞ {"x": 5, "y": 6}

<details>
  <summary><strong>Click here for a hint!</strong></summary>
  Loop through both dictionaries. If the key is in both, make a list of the two values.
</details>

<!-- grade_id: cell-e6d5653be2e068cc, points: 0, grade: false, solution: true -->
```python
def smart_merge(d1, d2):
    result = d1.copy()

    for key, value in d2.items():
        if key in result:
            result[key] = [result[key], value]
        else:
            result[key] = value

    return result
```

<!-- grade_id: cell-aa3483b556b16e29, points: 1, grade: true, solution: false -->
```python
# Tests
assert smart_merge({"a": 1}, {"a": 2}) == {"a": [1, 2]}
assert smart_merge({"x": 5}, {"y": 6}) == {"x": 5, "y": 6}
assert smart_merge({"a": 1, "b": 2}, {"b": 3, "c": 4}) == {"a": 1, "b": [2, 3], "c": 4}
```

## Exercise 6: Most Frequent Element

**Task**: Write a function `most_frequent(lst)` that returns the **element that appears most frequently** in the list `lst`.

📌 Use `collections.Counter`.

<details>
  <summary><strong>Click here for a hint!</strong></summary>
  Create a Counter and use <code>.most_common(1)</code> to get the most frequent item.
</details>

<!-- grade_id: cell-7c6f741391c9b533, points: 0, grade: false, solution: true -->
```python
from collections import Counter

def most_frequent(lst):
    counts = Counter(lst)
    return counts.most_common(1)[0][0]
```

<!-- grade_id: cell-eadde60cca3a82b3, points: 1, grade: true, solution: false -->
```python
# Tests
assert most_frequent([1, 2, 2, 3]) == 2
assert most_frequent(["a", "b", "a", "c", "a"]) == "a"
```

## Exercise 7: Group by First Letter

**Task**: Write a function `group_by_first_letter(words)` that takes a list of strings and returns a dictionary that groups the words by their **first letter**.

📌 Use `collections.defaultdict(list)`.

<details>
  <summary><strong>Click here for a hint!</strong></summary>
  Use <code>defaultdict(list)</code> and loop over each word using <code>word[0]</code> as the key.
</details>

<!-- grade_id: cell-2aa67e0b44e033d6, points: 0, grade: false, solution: true -->
```python
from collections import defaultdict

def group_by_first_letter(words):
    groups = defaultdict(list)
    for word in words:
        first_letter = word[0]
        groups[first_letter].append(word)

    return groups
```

<!-- grade_id: cell-413f7effe7fc411a, points: 1, grade: true, solution: false -->
```python
# Tests
assert group_by_first_letter(["apple", "banana", "apricot"]) == {"a": ["apple", "apricot"], "b": ["banana"]}
```

## Exercise 8: Rotate a List

**Task**: Write a function `rotate_list(lst, k)` that rotates the list **right** by `k` steps.

📌 Use `collections.deque`.

Example:
```python
rotate_list([1, 2, 3, 4], 1) => [4, 1, 2, 3]

<details>
  <summary><strong>Click here for a hint!</strong></summary>
  Convert to a <code>deque</code>, use <code>rotate(k)</code>, then convert back to a list.
</details>

<!-- grade_id: cell-9babe072b7bfa8a9, points: 0, grade: false, solution: true -->
```python
from collections import deque

def rotate_list(lst, k):
    d = deque(lst)
    d.rotate(k)
    return list(d)
```

<!-- grade_id: cell-6204acfa6999ca41, points: 1, grade: true, solution: false -->
```python
# Tests
assert rotate_list([1, 2, 3, 4], 1) == [4, 1, 2, 3]
assert rotate_list([1, 2, 3], 2) == [2, 3, 1]
```

## Exercise 9: Flatten a Nested List

**Task**: Write a function `flatten(lst)` that takes a **list of lists** and returns a **single flattened list**.

Example:
```python
flatten([[1, 2], [3, 4]]) => [1, 2, 3, 4]

<details>
  <summary><strong>Click here for a hint!</strong></summary>
  Use a nested loop or a list comprehension: <code>[x for sublist in lst for x in sublist]</code>
</details>

<!-- grade_id: cell-d6580fe8767b2124, points: 0, grade: false, solution: true -->
```python
def flatten(lst):
    flat_list = []
    for sublist in lst:
        for item in sublist:
            flat_list.append(item)
    return flat_list
```

<!-- grade_id: cell-f3fac42d40535987, points: 1, grade: true, solution: false -->
```python
# Tests
assert flatten([[1, 2], [3, 4]]) == [1, 2, 3, 4]
assert flatten([[], [5], [6, 7]]) == [5, 6, 7]
```

## Exercise 10: Word Frequency Counter

**Task**: Write a function `word_frequency(sentence)` that takes a string and returns a dictionary with the **frequency of each word**.

📌 Ignore punctuation. Treat words as lowercase.

Example:
```python
word_frequency("Hello hello world!") ➞ {"hello": 2, "world": 1}

<details>
  <summary><strong>Click here for a hint!</strong></summary>
  Use <code>.lower()</code>, <code>.split()</code>, and <code>collections.Counter</code>.
</details>

<!-- grade_id: cell-29f0f65e8656992d, points: 0, grade: false, solution: true -->
```python
from collections import Counter
import string
def word_frequency(sentence):
    sentence = sentence.lower()
    clean_sentence = sentence.translate(str.maketrans('', '', string.punctuation))
    words = clean_sentence.split()
    return dict(Counter(words))
```

<!-- grade_id: cell-c88b754eb51307a1, points: 1, grade: true, solution: false -->
```python
# Tests
assert word_frequency("One fish two fish red fish blue fish") == {
    "one": 1, "fish": 4, "two": 1, "red": 1, "blue": 1
}
assert word_frequency("Hello hello world!") == {"hello": 2, "world": 1}
```

## Optional Exercise 1: Leaderboard from Scores

**Task**: Write a function `top_scores(records)` where `records` is a **list of (name, score)** tuples.

Return a list of names sorted from **highest to lowest score**.

📌 If scores are equal, preserve the original order.

<details>
  <summary><strong>Click here for a hint!</strong></summary>
  Use <code>sorted()</code> with <code>key=lambda x: -x[1]</code> or <code>reverse=True</code>. Keep only the names.
</details>

<!-- grade_id: cell-ca2a52736277f9aa, points: 0, grade: false, solution: true -->
```python
def top_scores(records):
    sorted_records = sorted(records, key=lambda x: x[1], reverse=True)
    return [name for name, score in sorted_records]
```

<!-- grade_id: cell-cf73dbff5903046f, points: 1, grade: true, solution: false -->
```python
# Tests
data = [("Alice", 80), ("Bob", 95), ("Charlie", 90)]
assert top_scores(data) == ["Bob", "Charlie", "Alice"]
```

## Optional Exercise 2: Update Nested Inventory

**Task**: Write a function `update_inventory(inventory, category, item, quantity)`.

The `inventory` is a dictionary where each key is a category, and each value is another dictionary of `item: quantity`.

Update the quantity of the `item` under the given `category`.

📌 If the category or item doesn’t exist, create them.

<details>
  <summary><strong>Click here for a hint!</strong></summary>
  Use <code>setdefault()</code> or check with <code>if not in</code> for both levels.
</details>

<!-- grade_id: cell-f5afd87ecf856d1c, points: 0, grade: false, solution: true -->
```python
def update_inventory(inventory, category, item, quantity):
    category_dict = inventory.setdefault(category, {})
    category_dict[item] = quantity

    return inventory
```

<!-- grade_id: cell-13c43ccf3a387984, points: 1, grade: true, solution: false -->
```python
# Tests
inv = {"fruits": {"apple": 10}}
assert update_inventory(inv, "fruits", "banana", 5) == {"fruits": {"apple": 10, "banana": 5}}
assert update_inventory(inv, "veggies", "carrot", 7) == {
    "fruits": {"apple": 10, "banana": 5},
    "veggies": {"carrot": 7}
}
```

## Optional Exercise 3: Longest Word Per Starting Letter

**Task**: Write a function `longest_words(words)` that returns a dictionary where:

- Each key is the **first letter** of a word,
- Each value is the **longest word** in the list that starts with that letter.

📌 If multiple words have the same length, return the first one.

<details>
  <summary><strong>Click here for a hint!</strong></summary>
  Use a loop and compare lengths using <code>len(word)</code>.
</details>

<!-- grade_id: cell-10a26150682f3264, points: 0, grade: false, solution: true -->
```python
def longest_words(words):
    result = {}
    for word in words:
        first_letter = word[0]
        if first_letter not in result:
            result[first_letter] = word
        else:
            if len(word) > len(result[first_letter]):
                result[first_letter] = word

    return result
```

<!-- grade_id: cell-e7aa5ab0e885ed04, points: 1, grade: true, solution: false -->
```python
# Tests
assert longest_words(["apple", "ant", "banana", "boat", "cat"]) == {
    "a": "apple", "b": "banana", "c": "cat"
}
```