# 🧱 Python Practice: Functions, Parameters, and Recursion

Welcome to this hands-on practice notebook on **Functions, Parameters, and Recursion**! This notebook is designed to help you build a strong foundation in Python's function mechanics, preparing you for more advanced topics like functional programming and machine learning.

The exercises progress from basic syntax to moderately complex use cases such as using functions as arguments, filtering, mapping, conditional logic, and recursive thinking. By completing them, you'll develop the confidence to write expressive and efficient Python code.

## 🧭 Instructions
- This notebook contains **10 mandatory exercises** and **3 optional exercises** that are more complex.
- Each task is focused on key Python topics, including argument types, lambda functions, and recursion.
- Use the provided hints if you get stuck.
- You can run the tests as many times as needed.
- Total estimated time: **60 minutes**.
- Practice, think critically, and enjoy the process!

## 🧠 Concepts Covered
- Defining and calling functions
- Positional, keyword, default, and variable-length arguments (`*args`, `**kwargs`)
- Using built-in functions like `map`, `filter`, and `reduce`
- Functions as first-class objects (e.g., passing functions as arguments)
- Conditional transformations using lambda expressions
- Writing recursive solutions for common problems
- Realistic string and numeric manipulation

---

Let's dive in!

## Exercise 1: Dynamic Greeting

**Task**: Write a function `greet_user` that takes a name (string) and an hour (integer between 0 and 23) and returns a personalized greeting using an f-string. The greeting should change depending on the time of day:

- From 5 to 11 (inclusive): "Good morning, [name]! Welcome to Python."
- From 12 to 17 (inclusive): "Good afternoon, [name]! Welcome to Python."
- From 18 to 21 (inclusive): "Good evening, [name]! Welcome to Python."
- Otherwise: "Hello, [name]! Welcome to Python."

If the hour is outside the 0-23 range, return `"Invalid hour."`

<details>
  <summary><strong>Click here for a hint!</strong></summary>
  Use conditional statements to check the hour and select the greeting. Use an f-string to insert the name in the greeting. Don't forget to handle invalid hours.
</details>

<!-- grade_id: cell-e5d62626b4da425c, points: 0, grade: false, solution: true -->
```python
def greet_user(name, hour):
    if hour < 0 or hour > 23:
        return "Invalid hour."
    
    if 5 <= hour <= 11:
        greeting = "Good morning"
    elif 12 <= hour <= 17:
        greeting = "Good afternoon"
    elif 18 <= hour <= 21:
        greeting = "Good evening"
    else:
        return f"Hello, {name}! Welcome to Python."
        
    return f"{greeting}, {name}! Welcome to Python."
```

<!-- grade_id: cell-96f1f4b098135ed0, points: 1, grade: true, solution: false -->
```python
# Tests
assert greet_user("Alice", 9) == "Good morning, Alice! Welcome to Python.", "Morning greeting failed."
assert greet_user("Bob", 14) == "Good afternoon, Bob! Welcome to Python.", "Afternoon greeting failed."
assert greet_user("Eve", 19) == "Good evening, Eve! Welcome to Python.", "Evening greeting failed."
assert greet_user("Sam", 23) == "Hello, Sam! Welcome to Python.", "Night greeting failed."
assert greet_user("", 10) == "Good morning, ! Welcome to Python.", "Empty name handling failed."
assert greet_user("Tom", -1) == "Invalid hour.", "Invalid hour negative test failed."
assert greet_user("Tom", 24) == "Invalid hour.", "Invalid hour upper bound test failed."
```

## Exercise 2: Function with Default Parameters and Shape Selection

**Task**: Write a function `calculate_area` that calculates the area of a shape. The function takes the following parameters:

- `shape` (string): either `"rectangle"` or `"circle"`.
- If `shape` is `"rectangle"`, it takes `length` (required) and `width` (optional, default 5).
- If `shape` is `"circle"`, it takes `radius` (required).

The function should return the area of the specified shape. If parameters are missing or invalid (e.g., negative or zero), return `"Invalid dimensions"`.

Use `pi = 3.1416` for the circle area calculation.

<details>
  <summary><strong>Click here for a hint!</strong></summary>
  Use conditional statements to check the shape. For the rectangle, multiply length and width. For the circle, use <code>pi * radius²</code>. Validate inputs for positive values.
</details>

<!-- grade_id: cell-5451beca592bec33, points: 0, grade: false, solution: true -->
```python
def calculate_area(shape, length=None, width=5, radius=None):
    if shape == "circle":
        if radius is None or radius <= 0:
            return "Invalid dimensions"
        return 3.1416 * radius ** 2
    
    elif shape == "rectangle":
        if length is None or length <= 0 or width <= 0:
            return "Invalid dimensions"
        return length * width
    
    else:
        return "Invalid shape"
```

<!-- grade_id: cell-8b6c261274474fff, points: 1, grade: true, solution: false -->
```python
# Tests
assert calculate_area("rectangle", length=10) == 50
assert calculate_area("rectangle", length=10, width=2) == 20
assert calculate_area("circle", radius=3) == 28.2744
assert calculate_area("circle", radius=-3) == "Invalid dimensions"
assert calculate_area("rectangle", length=0, width=5) == "Invalid dimensions"
assert calculate_area("triangle", length=10, width=5) == "Invalid shape"
```

## Exercise 3: Mean and Median with *args

**Task**: Write a function `mean_and_median` that accepts any number of numeric arguments (`*args`) and returns a formatted string showing the mean and median of the numbers, e.g., `"Mean: 4.5, Median: 5"`. If no arguments are passed, return `"No numbers provided"`.

You can assume the inputs are numeric. Use the built-in `sorted()` function to help calculate the median.

<details>
  <summary><strong>Click here for a hint!</strong></summary>
  <ul>
    <li><strong>Mean</strong> is the sum divided by the count.</li>
    <li><strong>Median</strong> is the middle value in the sorted list (or the average of the two middle values if the count is even).</li>
    <li>Handle the <code>empty input</code> case separately.</li>
  </ul>
</details>

1 4 10 **15** 21 23 29 => median: 15

2 **15 48** 90 => median: (15 + 48) / 2 = 31.5


<!-- grade_id: cell-464305e2c8bc61df, points: 0, grade: false, solution: true -->
```python
def mean_and_median(*args):
    if not args:
        return "No numbers provided" 

    sorted_args = sorted(args)
    n = len(sorted_args)
    
    mean = sum(sorted_args) / n
    
    if n % 2 == 1:
        median = sorted_args[n // 2]
    else:
        median = (sorted_args[n // 2 - 1] + sorted_args[n // 2]) / 2
        
    return f"Mean: {mean}, Median: {median}"
```

<!-- grade_id: cell-1d3f019da8b7be82, points: 1, grade: true, solution: false -->
```python
# Tests
assert mean_and_median(1, 2, 3, 4, 5) == "Mean: 3.0, Median: 3"
assert mean_and_median(1, 2, 3, 4) == "Mean: 2.5, Median: 2.5"
assert mean_and_median() == "No numbers provided"
assert mean_and_median(7) == "Mean: 7.0, Median: 7"
```

## Exercise 4: Calculate Total Cost with Discounts by Item Index

**Task**: Write a function `calculate_total` that takes any number of item prices as positional arguments (`*args`) and any number of keyword arguments (`**kwargs`) representing discounts (percentages) on specific items, where the keys are **string indices** of the items (starting from `"0"`).

- Convert the string keys to integers to identify which item to discount.
- Apply the discount percentage to the corresponding item's price.
- Return the total cost after applying discounts.
- If no items are passed, return 0.

<details>
  <summary><strong>Click here for a hint!</strong></summary>
    Use <code>enumerate</code> on the prices and convert <code>kwargs</code> keys from strings to integers to apply discounts.
</details>

<!-- grade_id: cell-00234848a37e852c, points: 0, grade: false, solution: true -->
```python
def calculate_total(*args, **kwargs):
    if not args:
        return 0
    
    total = 0
    
    for i, price in enumerate(args):
        index_str = str(i)
        
        if index_str in kwargs:
            discount_percent = kwargs[index_str]
            total += price * (1 - discount_percent / 100)
        else:
            total += price
            
    return total
```

<!-- grade_id: cell-d9a97e39ccba1f82, points: 1, grade: true, solution: false -->
```python
# Tests
assert abs(calculate_total(100, 200, 300) - 600) < 1e-6  # no discounts
assert abs(calculate_total(100, 200, 300, **{"1": 10}) - 580) < 1e-6  # 10% off second item
assert abs(calculate_total(50, 50, 50, **{"0": 50, "2": 20}) - 115) < 1e-6  # discounts on first and third
assert calculate_total() == 0

```

## Exercise 5: Apply Discount Using `map` and `lambda`

**Task**: Write a function `apply_discount` that takes a list of prices and a discount percentage, and returns a new list of prices after applying the discount to each price.

- Use the `map` function combined with a `lambda` expression to apply the discount.
- The discount is a percentage (e.g., 10 means 10% off).
- Return the discounted prices as a list of floats.

<details>
  <summary><strong>Click here for a hint!</strong></summary>
    Use <code>map</code> with a lambda that multiplies each price by <code>(1 - discount / 100)</code>.
</details>

<!-- grade_id: cell-ff6b0c709e960128, points: 0, grade: false, solution: true -->
```python
def apply_discount(prices, discount):
    if not (0 <= discount <= 100):
        return "Invalid discount percentage."
    
    discounted_prices = list(map(lambda p: float(p * (1 - discount / 100)), prices))
    
    return discounted_prices
```

<!-- grade_id: cell-2ff6de4081447055, points: 1, grade: true, solution: false -->
```python
# Tests
assert apply_discount([100, 200, 300], 10) == [90.0, 180.0, 270.0]
assert apply_discount([50, 50, 50], 50) == [25.0, 25.0, 25.0]
assert apply_discount([], 20) == []
# Hidden tests
```

## Exercise 6: Filter Even Numbers Greater Than Threshold

**Task**: Write a function `filter_even_above` that takes a list of integers and a threshold value, and returns a new list containing only the even numbers strictly greater than the threshold.

- Use the `filter` function combined with a `lambda` to check both conditions.
- Return the filtered list sorted in ascending order.

<details>
  <summary><strong>Click here for a hint!</strong></summary>
    Use <code>filter</code> with a lambda that checks if a number is even (<code>x % 2 == 0</code>) and greater than the threshold.
</details>

<!-- grade_id: cell-4bbbcc49670c8d38, points: 0, grade: false, solution: true -->
```python
def filter_even_above(numbers, threshold):
    filtered_iterator = filter(lambda num: num % 2 == 0 and num > threshold, numbers)
    
    return sorted(list(filtered_iterator))
```

<!-- grade_id: cell-35c501b1cd258dff, points: 1, grade: true, solution: false -->
```python
# Tests
assert filter_even_above([1, 2, 3, 4, 5, 6], 3) == [4, 6]
assert filter_even_above([10, 13, 15, 20, 22], 15) == [20, 22]
assert filter_even_above([2, 4, 6], 6) == []
assert filter_even_above([], 0) == []
```

## Exercise 7: Calculate Product of a List Using `reduce`

**Task**: Write a function `product_of_list` that takes a list of numbers and returns the product of all elements.

- Use `functools.reduce` with a `lambda` function to multiply the elements.
- If the list is empty, return 1.


<details>
  <summary><strong>Click here for a hint!</strong></summary>
    Import <code>reduce</code> from <code>functools</code>. Use it with a lambda that multiplies two numbers.
</details>

<!-- grade_id: cell-275a57c676cced60, points: 0, grade: false, solution: true -->
```python
from functools import reduce

def product_of_list(numbers):
    if not numbers:
        return 1
    return reduce(lambda x, y: x * y, numbers)
```

<!-- grade_id: cell-38da2d29fa09aaf9, points: 1, grade: true, solution: false -->
```python
# Tests
assert product_of_list([1, 2, 3, 4]) == 24
assert product_of_list([5]) == 5
assert product_of_list([]) == 1
assert product_of_list([10, 0, 2]) == 0
```

## Exercise 8: Recursive Factorial for Multiple Numbers

**Task**: Write a recursive function `factorial` that computes the factorial of a single non-negative integer.

Then, write another function `factorials` that accepts any number of non-negative integers as positional arguments (`*args`) and returns a list of their factorials using the recursive `factorial` function.

- Assume the input integers are all ≥ 0.
- Use recursion for the factorial calculation.

<details>
  <summary><strong>Click here for a hint!</strong></summary>
  The factorial of 0 is 1. For n > 0, factorial(n) = n * factorial(n-1). Use a loop or <code>map</code> to apply factorial to all <code>*args</code>.
</details>

<!-- grade_id: cell-10c9a8e9d7a6bfbb, points: 0, grade: false, solution: true -->
```python
def factorial(n):
    if n == 0 or n == 1:
        return 1
    
    else:
        return n * factorial(n - 1)

def factorials(*args):
    return [factorial(n) for n in args]
```

<!-- grade_id: cell-54c09cb9af736244, points: 1, grade: true, solution: false -->
```python
# Tests
assert factorials(0, 1, 3, 5) == [1, 1, 6, 120]
assert factorials(4, 6) == [24, 720]
assert factorials() == []
```

## Exercise 9: Format User Info with `**kwargs`

**Task**: Write a function `format_user_info` that accepts any number of keyword arguments representing user info (all values are strings) and returns a single formatted string combining the info.

- The output should be formatted as `"key1: value1, key2: value2, ..."`, sorted by keys alphabetically.
- If no info is provided, return the string `"No user info provided"`.

<details>
  <summary><strong>Click here for a hint!</strong></summary>
    Use <code>kwargs.items()</code> and <code>sorted()</code> to sort keys. Use a list comprehension or <code>join</code> to build the string.
</details>

<!-- grade_id: cell-c1d50157c54731d0, points: 0, grade: false, solution: true -->
```python
def format_user_info(**kwargs):
    if not kwargs:
        return "No user info provided"
    
    sorted_keys = sorted(kwargs.keys())
    
    parts = []
    for key in sorted_keys:
        parts.append(f"{key}: {kwargs[key]}")
    
    return ", ".join(parts)
```

<!-- grade_id: cell-f782c238018f6f00, points: 1, grade: true, solution: false -->
```python
# Tests
assert format_user_info(name="Alice", city="Paris") == "city: Paris, name: Alice"
assert format_user_info(age="30", job="Engineer") == "age: 30, job: Engineer"
assert format_user_info() == "No user info provided"
```

## Exercise 10: Recursive String Reversal

**Task**: Write a recursive function `reverse_string` that takes a string and returns its reversal.

- Do not use Python slicing or built-in reversed functions.
- Base case: empty string returns empty string.

<details>
  <summary><strong>Click here for a hint!</strong></summary>
    Use recursion: <code>reverse_string(s) = last_char + reverse_string(rest_of_string)</code>.
</details>

<!-- grade_id: cell-bd03bf2571d0f076, points: 0, grade: false, solution: true -->
```python
def reverse_string(s):
    if s == "":
        return ""
    
    return s[-1] + reverse_string(s[:-1])
```

<!-- grade_id: cell-02b1467e65259967, points: 1, grade: true, solution: false -->
```python
# Tests
assert reverse_string("hello") == "olleh"
assert reverse_string("") == ""
assert reverse_string("a") == "a"
assert reverse_string("racecar") == "racecar"
```

## Optional Exercise 1: Deep Sum

**Task**: Write a recursive function `deep_sum` that takes a list which may contain integers or other lists (possibly nested multiple levels) and returns the total sum of all integers.


<details>
  <summary><strong>Click here for a hint!</strong></summary>
    Use <code>isinstance(x, list)</code> to check if an element is a list, and recurse accordingly.
</details>

<!-- grade_id: cell-021fd3af10958cc5, points: 0, grade: false, solution: true -->
```python
def deep_sum(lst):
    total = 0
    for item in lst:
        if isinstance(item, list):
            total += deep_sum(item)
        elif isinstance(item, (int, float)):
            total += item
    return total
```

<!-- grade_id: cell-9ae760a90b197033, points: 1, grade: true, solution: false -->
```python
# Tests
assert deep_sum([1, 2, [3, 4], 5]) == 15
assert deep_sum([1, [2, [3, [4]]]]) == 10
assert deep_sum([]) == 0
```

## Optional Exercise 2: Fibonacci with Memoization

**Task**: Write a recursive function `fib_memo` that returns the nth Fibonacci number (0-based), using a dictionary to memoize intermediate results for efficiency.

- Use a default parameter `memo={}` to store calculated Fibonacci numbers.
- `fib_memo(0) = 0`, `fib_memo(1) = 1`.

<details>
  <summary><strong>Click here for a hint!</strong></summary>
    Check if <code>n</code> is in <code>memo</code>. If not, calculate recursively and store it.
</details>

<!-- grade_id: cell-5235eff0fed5c37f, points: 0, grade: false, solution: true -->
```python
def fib_memo(n, memo={}):
    if n in memo:
        return memo[n]
    if n <= 1:
        return n
    memo[n] = fib_memo(n - 1, memo) + fib_memo(n - 2, memo)
    return memo[n]
```

<!-- grade_id: cell-b563164a9204e810, points: 1, grade: true, solution: false -->
```python
# Tests
assert fib_memo(0) == 0
assert fib_memo(1) == 1
assert fib_memo(5) == 5
assert fib_memo(10) == 55
```

## Optional Exercise 3: Conditional Apply with Function Argument

**Task**: Write a function `conditional_apply` that takes a list of items, a function `func`, and a condition function `condition`.

Return a new list where `func(x)` is applied **only to those elements** for which `condition(x)` is `True`. Leave all other elements unchanged.

<details>
  <summary><strong>Click here for a hint!</strong></summary>
  Use a list comprehension. For each <code>x</code>, check if <code>condition(x)</code> is true — if yes, apply <code>func(x)</code>, otherwise just keep <code>x</code>.
</details>

<!-- grade_id: cell-d0e5da550a8c7d21, points: 0, grade: false, solution: true -->
```python
def conditional_apply(items, func, condition):
    return [func(item) if condition(item) else item for item in items]
```

<!-- grade_id: cell-f0328f6f7e6fdd23, points: 1, grade: true, solution: false -->
```python
# Tests
assert conditional_apply([1, 2, 3, 4], lambda x: x * 10, lambda x: x % 2 == 0) == [1, 20, 3, 40]
assert conditional_apply(["hello", "world", "hi"], str.upper, lambda x: len(x) > 3) == ["HELLO", "WORLD", "hi"]
assert conditional_apply([10, -5, 0], abs, lambda x: x < 0) == [10, 5, 0]
```