# 🧱 Python Practice: Pandas

Welcome to this hands-on practice notebook on **Pandas**, Python’s powerful data analysis library. This notebook is designed to simulate a real-world data scenario where you’ll work with **multiple CSV files** representing students, courses, and grades.

You’ll practice tasks such as **merging DataFrames**, **grouping and aggregating data**, **filtering with conditions**, and **analyzing trends** across datasets — all essential skills for data analysis and machine learning projects.

As you progress, you’ll explore how to reshape, summarize, and gain insights from data in ways that mirror real tasks analysts and scientists do in industry.

---

## 🧭 Instructions
- This notebook contains **10 mandatory exercises** and **3 optional bonus exercises** that are slightly more advanced.
- Each task builds on the same data model, so keep your CSVs in the same folder as the notebook.
- Use the provided hints if you get stuck.
- You can run the tests as many times as needed.
- Total estimated time: **60 minutes**.
- Focus on clarity, logic, and learning to think like a data analyst!

---

## 🧠 Concepts Covered
- Reading CSV files with `pd.read_csv()`
- Merging/joining multiple DataFrames
- Filtering rows using conditions
- Grouping and aggregating data with `groupby()`
- Sorting and reshaping DataFrames
- Analyzing numeric columns
- Performing cumulative operations (`expanding()`, `cumsum()`)
- Writing reusable data functions

---

Let’s dive in!

## Exercise 1: Load Students CSV

**Task**: Write a function `load_students()` that reads the file `students.csv` and returns it as a DataFrame.

<details>
  <summary><strong>Click here for a hint!</strong></summary>
  Use <code>pd.read_csv("students.csv")</code> and make sure to return the DataFrame.
</details>

<!-- grade_id: cell-526dd27d89987105, points: 0, grade: false, solution: true -->
```python
import pandas as pd

def load_students():
    return pd.read_csv("students.csv")
```

<!-- grade_id: cell-8fe39ed5d93a437e, points: 1, grade: true, solution: false -->
```python
# Tests
df = load_students()
assert isinstance(df, pd.DataFrame)
assert "student_id" in df.columns
assert df.shape[0] == 10
```

## Exercise 2: Merge All CSVs

**Task**: Write a function `merge_all()` that:

1. Loads `students.csv`, `courses.csv`, and `grades.csv`.
2. Merges them into a single DataFrame using `student_id` |and `course_id`.
3. Returns the merged DataFrame.

<details>
  <summary><strong>Click here for a hint!</strong></summary>
  Use <code>pd.merge</code> twice: first on <code>student_id</code>, then on <code>course_id</code>.
</details>

<!-- grade_id: cell-81044af7ce9a7736, points: 0, grade: false, solution: true -->
```python
def merge_all():
    students = pd.read_csv("students.csv")
    courses = pd.read_csv("courses.csv")
    grades = pd.read_csv("grades.csv")
    
    merged = pd.merge(grades, students, on="student_id")
    merged = pd.merge(merged, courses, on="course_id")
    
    return merged
```

<!-- grade_id: cell-1451c4ef6d0e1210, points: 1, grade: true, solution: false -->
```python
# Tests
df = merge_all()
assert "student_id" in df.columns
assert "grade" in df.columns
assert "course_name" in df.columns
assert df.shape[0] >= 30
```

## Exercise 3: Average Grade per Student

**Task**: Write a function `average_grade_per_student()` that returns a DataFrame with one row per student and their average grade across all courses.

<details>
  <summary><strong>Click here for a hint!</strong></summary>
  Use <code>groupby</code> on <code>student_name</code> and apply <code>mean()</code> on <code>grade</code>.
</details>

<!-- grade_id: cell-f6c794ce2db94674, points: 0, grade: false, solution: true -->
```python
def average_grade_per_student():
    df = merge_all()
    result = df.groupby("student_name")["grade"].mean().reset_index()
    result.columns = ["student_name", "avg_grade"]
    return result
```

<!-- grade_id: cell-178ce36afa9a5c22, points: 1, grade: true, solution: false -->
```python
# Tests
df = average_grade_per_student()
assert "student_name" in df.columns
assert "avg_grade" in df.columns
assert df.shape[0] == 10
```

## Exercise 4: Average Grade per Course

**Task**: Write a function `average_grade_per_course()` that returns a DataFrame showing the average grade for each course.

<details>
  <summary><strong>Click here for a hint!</strong></summary>
  Group by <code>course_name</code> and apply <code>mean()</code> to the <code>grade</code> column.
</details>

<!-- grade_id: cell-dee75d25a78cb443, points: 0, grade: false, solution: true -->
```python
def average_grade_per_course():
    df = merge_all()
    result = df.groupby("course_name")["grade"].mean().reset_index()
    result.columns = ["course_name", "avg_grade"]
    return result
```

<!-- grade_id: cell-4c8bc1b686a7d571, points: 1, grade: true, solution: false -->
```python
# Tests
df = average_grade_per_course()
assert "course_name" in df.columns
assert "avg_grade" in df.columns
assert df.shape[0] == 6
```

## Exercise 5: Highest Grade per Course

**Task**: Write a function `highest_grade_per_course()` that returns a DataFrame with the highest grade for each course.

<details>
  <summary><strong>Click here for a hint!</strong></summary>
  Use <code>groupby</code> on <code>course_name</code> and apply <code>max()</code> to <code>grade</code>.
</details>

<!-- grade_id: cell-ebf9678a208cd27c, points: 0, grade: false, solution: true -->
```python
def highest_grade_per_course():
    df = merge_all()
    result = df.groupby("course_name")["grade"].max().reset_index()
    result.columns = ["course_name", "max_grade"]
    return result
```

<!-- grade_id: cell-7926ff9909591d9a, points: 1, grade: true, solution: false -->
```python
# Tests
df = highest_grade_per_course()
assert "max_grade" in df.columns
assert df["max_grade"].between(70, 100).all()
```

## Exercise 6: Gender Distribution by Course

**Task**: Write a function `gender_distribution()` that returns a table (DataFrame) showing how many male and female students are enrolled in each course.

<details>
  <summary><strong>Click here for a hint!</strong></summary>
  Use <code>groupby</code> with <code>size()</code> and <code>unstack()</code> to pivot the gender counts.
</details>

<!-- grade_id: cell-a6aeb606156184e4, points: 0, grade: false, solution: true -->
```python
def gender_distribution():
    df = merge_all()
    result = df.groupby(["course_name", "gender"]).size().unstack(fill_value=0)
    return result
```

<!-- grade_id: cell-a310b43f1f952300, points: 1, grade: true, solution: false -->
```python
# Tests
df = gender_distribution()
assert "F" in df.columns and "M" in df.columns
assert df.shape[0] == 6
```

## Exercise 7: Students Who Scored Above 90

**Task**: Write a function `students_above_90()` that returns a DataFrame of all rows where the `grade` is greater than 90.

<details>
  <summary><strong>Click here for a hint!</strong></summary>
  Use boolean filtering: <code>df[df["grade"] > 90]</code>.
</details>

<!-- grade_id: cell-11a205f9523d91b8, points: 0, grade: false, solution: true -->
```python
def students_above_90():
    df = merge_all()
    return df[df["grade"] > 90]
```

<!-- grade_id: cell-a742ecc5fb3ad473, points: 1, grade: true, solution: false -->
```python
# Tests
df = students_above_90()
assert (df["grade"] > 90).all()
assert "student_id" in df.columns
```

## Exercise 8: Pivot Table — Average Grade by Gender and Course

**Task**: Write a function `pivot_grade_gender_course()` that returns a pivot table showing the average grade by `gender` and `course_name`.

<details>
  <summary><strong>Click here for a hint!</strong></summary>
  Use <code>pd.pivot_table()</code> with <code>index="course_name"</code>, <code>columns="gender"</code>, and <code>values="grade"</code>.
</details>

<!-- grade_id: cell-3e8948b39bd51743, points: 0, grade: false, solution: true -->
```python
def pivot_grade_gender_course():
    df = merge_all()
    return pd.pivot_table(df, index="course_name", columns="gender", values="grade", aggfunc="mean")
```

<!-- grade_id: cell-8d4a11e4453540a3, points: 1, grade: true, solution: false -->
```python
# Tests
pivot = pivot_grade_gender_course()
assert "F" in pivot.columns and "M" in pivot.columns
assert pivot.shape[0] == 6
```

## Exercise 9: Count Number of Courses Taken by Each Student

**Task**: Write a function `course_count_per_student()` that returns a Series with `student_name` as index and the number of grades (i.e., courses taken) per student.

<details>
  <summary><strong>Click here for a hint!</strong></summary>
  Use <code>value_counts()</code> on the <code>student_name</code> column.
</details>

<!-- grade_id: cell-093286c799c43da2, points: 0, grade: false, solution: true -->
```python
def course_count_per_student():
    df = merge_all()
    return df["student_name"].value_counts()
```

<!-- grade_id: cell-01b2c162f57ddc0a, points: 1, grade: true, solution: false -->
```python
# Tests
s = course_count_per_student()
assert isinstance(s, pd.Series)
assert s.index.isin(["Alice", "Bob"]).any()
```

## Exercise 10: Normalize the Grade Column

**Task**: Write a function `normalize_grades()` that returns a DataFrame with a new column `grade_norm` where each grade is normalized using z-score:  
`z_score = (grade - mean) / std`

<details>
  <summary><strong>Click here for a hint!</strong></summary>
  Compute mean and std of <code>grade</code>, then use <code>apply(lambda x: (x - mean) / std)</code>.
</details>

<!-- grade_id: cell-080a20d3fc5b91e7, points: 0, grade: false, solution: true -->
```python
def normalize_grades():
    df = merge_all()
    mean = df["grade"].mean()
    std = df["grade"].std()
    df["grade_norm"] = (df["grade"] - mean) / std
    return df
```

<!-- grade_id: cell-f499e12cbbab2b43, points: 1, grade: true, solution: false -->
```python
# Tests
df = normalize_grades()
assert "grade_norm" in df.columns
assert abs(df["grade_norm"].mean()) < 1e-6  # close to 0
```

## Optional Exercise 1: Courses Never Taken by Any Female Student

**Task**: Write a function `courses_no_females()` that returns a list of course names where no female students are enrolled.

<details>
  <summary><strong>Click here for a hint!</strong></summary>
  Filter for <code>gender="F"</code>, then check which <code>course_names</code> are missing from the full set.
</details>

<!-- grade_id: cell-cc26a5194c4f217a, points: 0, grade: false, solution: true -->
```python
def courses_no_females():
    df = merge_all()
    all_courses = set(df["course_name"].unique())
    female_courses = set(df[df["gender"] == "F"]["course_name"].unique())
    return list(all_courses - female_courses)
```

<!-- grade_id: cell-91814335d5582431, points: 1, grade: true, solution: false -->
```python
# Tests
lst = courses_no_females()
assert isinstance(lst, list)
```

## Optional Exercise 2: Add “Passed” Column

**Task**: Write a function `add_passed_column()` that adds a new column `passed` to the DataFrame, set to `True` if `grade >= 60`, else `False`.

<details>
  <summary><strong>Click here for a hint!</strong></summary>
  Use <code>df["grade"] >= 60</code> and assign the result to <code>df["passed"]</code>.
</details>

<!-- grade_id: cell-df9b6dd43fc5a303, points: 0, grade: false, solution: true -->
```python
def add_passed_column():
    df = merge_all()
    df["passed"] = df["grade"] >= 60
    return df
```

<!-- grade_id: cell-998c9327d07c63e5, points: 1, grade: true, solution: false -->
```python
# Tests
df = add_passed_column()
assert "passed" in df.columns
assert set(df["passed"].unique()).issubset({True, False})
```

## Optional Exercise 3: Grade Progress Per Student (Cumulative Mean)

**Task**: Write a function `grade_progress_per_student()` that returns a DataFrame showing how each student's average grade evolves over time, ordered by `course_name` alphabetically.

- Include columns: `student_name`, `course_name`, `grade`, `avg_so_far`
- The `avg_so_far` column should contain the cumulative average grade for each student, computed as they progress through the courses alphabetically.

<details>
  <summary><strong>Click here for a hint!</strong></summary>
  Sort the merged DataFrame by <code>student_name</code> and <code>course_name</code>, then use <code>groupby()</code> + <code>expanding().mean()</code>.
</details>

<!-- grade_id: cell-c5868b27e8250a98, points: 0, grade: false, solution: true -->
```python
def grade_progress_per_student():
    df = merge_all()
    df = df.sort_values(["student_name", "course_name"])
    df["avg_so_far"] = df.groupby("student_name")["grade"].expanding().mean().reset_index(level=0, drop=True)
    return df[["student_name", "course_name", "grade", "avg_so_far"]]
```

<!-- grade_id: cell-fe97173408edb408, points: 1, grade: true, solution: false -->
```python
# Tests
df = grade_progress_per_student()
assert all(col in df.columns for col in ["student_name", "course_name", "grade", "avg_so_far"])
assert df["avg_so_far"].notnull().all()
```