# 🧱 Python Practice: NumPy — Introduction to NumPy Library

Welcome to this practice notebook on **NumPy**, Python’s powerhouse library for numerical computing and data manipulation. This notebook is designed to deepen your understanding of how NumPy handles **arrays**, **vectorized operations**, **broadcasting**, and **aggregation**.

You’ll work through problems that go beyond basic array creation — such as computing local maxima, normalizing data, rotating matrices, and performing operations using masks and axis-specific logic. These are the types of operations that form the backbone of data science, image processing, and machine learning workflows.

---

## 🧭 Instructions
- This notebook contains **10 mandatory exercises** and **3 optional (bonus) exercises** that are a bit more challenging.
- Each task focuses on a key NumPy concept, such as slicing, broadcasting, logical masking, or reshaping.
- Hints are available if you get stuck — take your time and explore!
- You can run the tests as many times as needed.
- Total estimated time: **60 minutes**.
- Think like a data scientist: manipulate, transform, and analyze efficiently.

---

## 🧠 Concepts Covered
- Advanced array slicing and indexing
- Axis-based aggregation (row-wise, column-wise)
- Broadcasting rules and vectorized computation
- Conditional logic and masking
- Data normalization and statistical operations
- Multi-dimensional reshaping and transformation
- Writing clean, reusable NumPy functions

---

Let’s dive in!

## Exercise 1: Column-Wise Mean Normalization

**Task**: Write a function `normalize_columns(arr)` that:

1. Accepts a **2D NumPy array**.
2. Returns a new array where each column has been **mean-centered**:
   - Subtract the column's mean from each element in that column.

📌 This is useful in data preprocessing (e.g., in ML pipelines).

<details>
  <summary><strong>Click here for a hint!</strong></summary>
  Use <code>arr.mean(axis=0)</code> to get the column-wise mean. Then use broadcasting to subtract it.
</details>

<!-- grade_id: cell-ac2d90c384d20203, points: 0, grade: false, solution: true -->
```python
import numpy as np

def normalize_columns(arr):
    col_means = arr.mean(axis=0)
    normalized_arr = arr - col_means
    return normalized_arr
```

<!-- grade_id: cell-c1e458b22ad555f3, points: 1, grade: true, solution: false -->
```python
# Tests
a = np.array([[1, 2],
              [3, 4],
              [5, 6]])
norm = normalize_columns(a)
assert np.allclose(norm.mean(axis=0), [0, 0], atol=1e-6)
```

## Exercise 2: Conditional Element Swap

**Task**: Write a function `swap_if_larger(arr)` that:

1. Accepts a **2D NumPy array** of shape `(n, 2)` — 2 columns only.
2. Swaps elements in each row **if the first column value is greater** than the second.

Example:
```python
Input:  [[5, 2], [1, 9], [3, 3]]
Output: [[2, 5], [1, 9], [3, 3]]

<details>
  <summary><strong>Click here for a hint!</strong></summary>
  Use Boolean indexing on <code>arr[:, 0] > arr[:, 1]</code> and swap elements row-wise using slicing.
</details>

<!-- grade_id: cell-11173ac738a52c92, points: 0, grade: false, solution: true -->
```python
def swap_if_larger(arr):
    res = arr.copy()
    mask = res[:, 0] > res[:, 1]
    res[mask, 0], res[mask, 1] = res[mask, 1], res[mask, 0].copy()
    return res
```

<!-- grade_id: cell-cd91213b3446ef57, points: 1, grade: true, solution: false -->
```python
# Tests
a = np.array([[5, 2], [1, 9], [3, 3]])
expected = np.array([[2, 5], [1, 9], [3, 3]])
assert np.array_equal(swap_if_larger(a), expected)
```

## Exercise 3: Checkerboard Pattern Generator

**Task**: Write a function `create_checkerboard(n)` that:

1. Returns an `n x n` NumPy array.
2. The array should contain 0s and 1s in a checkerboard (chessboard) pattern, starting with 0 at the top-left.

Example (`n=4`):

```python
[[0, 1, 0, 1],
 [1, 0, 1, 0],
 [0, 1, 0, 1],
 [1, 0, 1, 0]]

<details>
  <summary><strong>Click here for a hint!</strong></summary>
  Use <code>np.indices()</code> to generate row and column index grids, and sum them.
</details>

<!-- grade_id: cell-704b4b9ce44edd18, points: 0, grade: false, solution: true -->
```python
def create_checkerboard(n):
    rows, cols = np.indices((n, n))
    return (rows + cols) % 2
```

<!-- grade_id: cell-ab944e10ab0cf635, points: 1, grade: true, solution: false -->
```python
# Tests
cb = create_checkerboard(4)
assert cb.shape == (4, 4)
assert cb[0, 0] == 0 and cb[0, 1] == 1
assert cb[1, 0] == 1 and cb[1, 1] == 0
```

## Exercise 4: Find Most Frequent Element

**Task**: Write a function `most_frequent_element(arr)` that:

1. Accepts a 1D NumPy array of integers.
2. Returns the integer that appears **most frequently** in the array.
3. If multiple values tie, return the smallest one.

Example:
```python
most_frequent_element(np.array([1, 2, 2, 3, 3])) → 2

<details>
  <summary><strong>Click here for a hint!</strong></summary>
  Use <code>np.unique(arr, return_counts=True)</code> to get frequencies. Then use <code>np.argmax</code>.
</details>

<!-- grade_id: cell-524654cdc5b29500, points: 0, grade: false, solution: true -->
```python
def most_frequent_element(arr):
    values, counts = np.unique(arr, return_counts=True)
    index = np.argmax(counts)
    return values[index]
```

<!-- grade_id: cell-6c1b9ca17759bfa9, points: 1, grade: true, solution: false -->
```python
# Tests
a = np.array([1, 2, 2, 3, 3])
assert most_frequent_element(a) == 2
```

## Exercise 5: Zero Border Padding

**Task**: Write a function `add_zero_border(arr)` that:

1. Takes a 2D NumPy array.
2. Returns a new array with a **1-cell-wide border of zeros** added around the original.

Example:
```python
Input:
[[1, 2],
 [3, 4]]

Output:
[[0, 0, 0, 0],
 [0, 1, 2, 0],
 [0, 3, 4, 0],
 [0, 0, 0, 0]]

<details>
  <summary><strong>Click here for a hint!</strong></summary>
  Use <code>np.pad</code> with <code>pad_width=1</code> and <code>constant_values=0</code>.
</details>

<!-- grade_id: cell-bc1b7fdb5e053eec, points: 0, grade: false, solution: true -->
```python
def add_zero_border(arr):
    return np.pad(arr, pad_width=1, mode='constant', constant_values=0)
```

<!-- grade_id: cell-a7db2da7b7266f8a, points: 1, grade: true, solution: false -->
```python
a = np.array([[1, 2], [3, 4]])
res = add_zero_border(a)
assert res.shape == (4, 4)
assert np.array_equal(res[1:3, 1:3], a)
assert np.all(res[0] == 0) and np.all(res[:, -1] == 0)
```

## Exercise 6: Row with Highest Sum

**Task**: Write a function `row_with_max_sum(arr)` that:

1. Takes a 2D NumPy array.
2. Returns the **index of the row** with the highest sum of elements.

If there is a tie, return the **first such row index**.

<details>
  <summary><strong>Click here for a hint!</strong></summary>
  Use <code>arr.sum(axis=1)</code> to compute row sums, and <code>np.argmax</code> to get the index of the largest.
</details>

<!-- grade_id: cell-6999156227929953, points: 0, grade: false, solution: true -->
```python
def row_with_max_sum(arr):
    row_sums = arr.sum(axis=1)
    return np.argmax(row_sums)
```

<!-- grade_id: cell-921717e2066baa46, points: 1, grade: true, solution: false -->
```python
a = np.array([[1, 2], [4, 0], [1, 1]])
assert row_with_max_sum(a) == 1
```

## Exercise 7: Extract Diagonal Differences

**Task**: Write a function `diagonal_difference(matrix)` that:

1. Accepts a square 2D NumPy array.
2. Returns the absolute difference between the sums of its two diagonals.

Main diagonal: top-left to bottom-right  
Secondary diagonal: top-right to bottom-left

<details>
  <summary><strong>Click here for a hint!</strong></summary>
  Use <code>np.trace()</code> for the main diagonal and <code>np.fliplr()</code> with <code>np.trace()</code> for the other.
</details>

<!-- grade_id: cell-643d57719aa14473, points: 0, grade: false, solution: true -->
```python
def diagonal_difference(matrix):
    main_diag_sum = np.trace(matrix)
    secondary_diag_sum = np.trace(np.fliplr(matrix))
    return abs(main_diag_sum - secondary_diag_sum)
```

<!-- grade_id: cell-914a103ea592af88, points: 1, grade: true, solution: false -->
```python
mat = np.array([[1, 2, 3],
                [4, 5, 6],
                [9, 8, 9]])
assert diagonal_difference(mat) == 2
```

## Exercise 8: Replace Outliers

**Task**: Write a function `replace_outliers(arr)` that:

1. Takes a 1D NumPy array.
2. Replaces any element that is more than **3 standard deviations** from the mean with the **median value**.

Return the updated array.

<details>
  <summary><strong>Click here for a hint!</strong></summary>
  Use <code>np.mean</code> and <code>np.std</code> to compute thresholds. Replace using a mask and <code>np.median</code>.
</details>

<!-- grade_id: cell-135e41ca5396e5cb, points: 0, grade: false, solution: true -->
```python
def replace_outliers(arr):
    mu = np.mean(arr)
    sd = np.std(arr)
    med = np.median(arr)
    
    outlier_mask = np.abs(arr - mu) > 3 * sd
    
    res = arr.copy()
    res[outlier_mask] = med
    return res
```

<!-- grade_id: cell-9ac62d4451a9bd50, points: 1, grade: true, solution: false -->
```python
a = np.array([10, 12, 14, 16, 20, 21, 23, 5, 7, 9, 200])
res = replace_outliers(a)
assert res[-1] != 200
assert np.isclose(res[-1], np.median(res), atol=1e-5)
```

## Exercise 9: Reshape and Tile

**Task**: Write a function `repeat_reshape(arr)` that:

1. Accepts a 1D NumPy array of length `n`.
2. Returns a 2D array of shape `(n, n)` where **each row is a copy of the original array**.

Example:
```python
repeat_reshape(np.array([1, 2])) → [[1, 2], [1, 2]]

<details>
  <summary><strong>Click here for a hint!</strong></summary>
  Use <code>np.tile</code> or <code>np.repeat</code> combined with reshaping.
</details>

<!-- grade_id: cell-83a6a8913649bb54, points: 0, grade: false, solution: true -->
```python
def repeat_reshape(arr):
    n = len(arr)
    return np.tile(arr, (n, 1))
```

<!-- grade_id: cell-b68125ff216e6ee4, points: 1, grade: true, solution: false -->
```python
a = np.array([1, 2, 3])
out = repeat_reshape(a)
assert out.shape == (3, 3)
assert np.all(out[0] == a)
```

## Exercise 10: Normalize to 0–1 Range

**Task**: Write a function `normalize_to_unit(arr)` that:

1. Accepts a NumPy array of any shape.
2. Returns a new array where all values are scaled to the range **0 to 1**, based on min-max normalization.

Formula:
`normalized = (arr - min) / (max - min)`

<details>
  <summary><strong>Click here for a hint!</strong></summary>
  Use <code>arr.min()</code> and <code>arr.max()</code>. Watch for division by zero when all elements are the same.
</details>

<!-- grade_id: cell-0601f39c83ad4895, points: 0, grade: false, solution: true -->
```python
def normalize_to_unit(arr):
    arr_min = arr.min()
    arr_max = arr.max()
    
    if arr_max == arr_min:
        return np.zeros_like(arr, dtype=float)
        
    return (arr - arr_min) / (arr_max - arr_min)
```

<!-- grade_id: cell-45c51bec27ad65f7, points: 1, grade: true, solution: false -->
```python
a = np.array([10, 20, 30])
res = normalize_to_unit(a)
assert res.min() == 0 and res.max() == 1
```

## Optional Exercise 1: Find Local Maxima

**Task**: Write a function `find_local_maxima(arr)` that:

1. Accepts a 1D NumPy array.
2. Returns a new array of the same length, containing:
   - `True` where the element is greater than **both** neighbors,
   - `False` elsewhere.

First and last elements should always be `False`.

Example:
```python
[1, 3, 2, 5, 1] → [False, True, False, True, False]

<details>
  <summary><strong>Click here for a hint!</strong></summary>
  Use slicing like <code>arr[1:-1]</code> and compare to neighbors with broadcasting.
</details>

<!-- grade_id: cell-7e87f9e614f2190c, points: 0, grade: false, solution: true -->
```python
def find_local_maxima(arr):
    res = np.zeros_like(arr, dtype=bool)
    if len(arr) < 3:
        return res
    
    mid = arr[1:-1]
    left = arr[:-2]
    right = arr[2:]
    
    res[1:-1] = (mid > left) & (mid > right)
    return res
```

<!-- grade_id: cell-a13d46f2a2ff487e, points: 1, grade: true, solution: false -->
```python
a = np.array([1, 3, 2, 5, 1])
out = find_local_maxima(a)
assert np.array_equal(out, [False, True, False, True, False])
```

## Optional Exercise 2: Rotate Matrix 90° Clockwise

**Task**: Write a function `rotate_90_clockwise(matrix)` that:

1. Accepts a 2D square NumPy array.
2. Returns a new array rotated 90 degrees clockwise.

Example:
```python
[[1, 2],
 [3, 4]] → [[3, 1],
            [4, 2]]

<details>
  <summary><strong>Click here for a hint!</strong></summary>
  First try <code>np.transpose()</code> or <code>.T</code> to flip rows/cols, then reverse rows.
</details>

<!-- grade_id: cell-0c7b8b9e6733646b, points: 0, grade: false, solution: true -->
```python
def rotate_90_clockwise(matrix):
    return np.flip(matrix.T, axis=1)
```

<!-- grade_id: cell-1be4ca44c38f7aad, points: 1, grade: true, solution: false -->
```python
m = np.array([[1, 2], [3, 4]])
r = rotate_90_clockwise(m)
expected = np.array([[3, 1], [4, 2]])
assert np.array_equal(r, expected)
```

## Optional Exercise 3: Row-wise Z-Score Normalization

**Task**: Write a function `row_z_scores(arr)` that:

1. Accepts a **2D NumPy array**.
2. Returns a new array where each row has been **standardized using Z-scores**:
   - Subtract the **mean of each row**
   - Divide by the **standard deviation of that row**

This is common in data preprocessing when features need to be normalized row-wise.

Example:
```python
Input:
[[1, 2, 3],
 [4, 5, 6]]

Output:
[[-1.224, 0.0, 1.224],
 [-1.224, 0.0, 1.224]]

<details>
  <summary><strong>Click here for a hint!</strong></summary>
  Use <code>arr.mean(axis=1, keepdims=True)</code> and <code>arr.std(axis=1, keepdims=True)</code> to make broadcasting work across rows.
</details>

<!-- grade_id: cell-fee66484a12acf98, points: 0, grade: false, solution: true -->
```python
def row_z_scores(arr):
    mu = arr.mean(axis=1, keepdims=True)
    sd = arr.std(axis=1, keepdims=True)
    return (arr - mu) / sd
```

<!-- grade_id: cell-30f1087bf3194ca0, points: 1, grade: true, solution: false -->
```python
a = np.array([[1, 2, 3], [4, 5, 6]])
res = row_z_scores(a)
assert res.shape == (2, 3)
assert np.allclose(res.mean(axis=1), [0, 0], atol=1e-6)
assert np.allclose(res.std(axis=1), [1, 1], atol=1e-6)
```