# 🧱 Python Practice: Object-Oriented Programming

Welcome to this hands-on practice notebook on **Object-Oriented Programming (OOP)** in Python! This notebook will help you build a solid foundation in how to **design and work with classes**, the cornerstone of OOP.

You’ll learn how to define your own types, add attributes and methods, use constructors, and apply key principles like encapsulation. Whether you're modeling a `Dog`, a `Car`, or a `BankAccount`, you'll practice structuring your code in a more modular and reusable way — just like real-world developers do.

---

## 🧭 Instructions
- This notebook contains **10 mandatory exercises** and **3 optional exercises** that are a bit more advanced.
- Each task focuses on a fundamental OOP concept such as writing a class, using constructors, or defining methods.
- Hints are provided if you get stuck — feel free to use them!
- You can run the tests as many times as needed.
- Total estimated time: **60 minutes**.
- Think in terms of objects, stay curious, and enjoy learning!

---

## 🧠 Concepts Covered
- Creating custom classes using `class`
- Adding attributes to objects
- Writing instance methods
- Using the `__init__` constructor
- Working with `self`
- Representing objects with `__str__` or `__repr__`
- Object instantiation and attribute access
- Basic encapsulation principles

---

Let’s dive in!

## Exercise 1: Define a `Book` Class with Attributes

**Task**: Write a class called `Book` that represents a book in a library. It should:

1. Be initialized with three attributes: `title` (string), `author` (string), and `pages` (integer).
2. Store each attribute as an instance variable.

For example, creating a book like this:

```python
b = Book("1984", "George Orwell", 328)
```
should allow access like:
```python
b.title → "1984"
b.pages → 328
```

## Exercise 10: Count Books by Title Keyword

**Task**: Add a method `count_books_with(keyword)` to the `Library` class that:

- Returns the number of books that contain the given keyword (case-insensitive) in their title.

<details>
  <summary><strong>Click here for a hint!</strong></summary>
  Use <code>__init__</code> to receive and assign <code>title</code>, <code>author</code>, and <code>pages</code> to <code>self</code>.
</details>

<!-- grade_id: cell-8017f1722ea566ad, points: 0, grade: false, solution: true -->
```python
class Book:
    def __init__(self, title, author, pages):
        self.title = title
        self.author = author
        self.pages = pages
```

<!-- grade_id: cell-004faebdf5b41591, points: 1, grade: true, solution: false -->
```python
# Tests
book = Book("1984", "George Orwell", 328)
assert book.title == "1984"
assert book.author == "George Orwell"
assert book.pages == 328
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In[1], line 2
      1 # Tests
----> 2 book = Book("1984", "George Orwell", 328)
      3 assert book.title == "1984"
      4 assert book.author == "George Orwell"

NameError: name 'Book' is not defined
```

## Exercise 2: Add a `summary` Method to the `Book` Class

**Task**: Extend your `Book` class to include a method `summary()` that:

- Returns a string in the format:  
  `"Title: <title> | Author: <author> | Pages: <pages>"`

Example:
```python
b = Book("1984", "George Orwell", 328)
b.summary()
# Output: "Title: 1984 | Author: George Orwell | Pages: 328"

<details>
  <summary><strong>Click here for a hint!</strong></summary>
  Define a method <code>summary(self)</code> and use an f-string to return the desired format using the object's attributes.
</details>

<!-- grade_id: cell-58cd5092c51ad888, points: 0, grade: false, solution: true -->
```python
class Book:
    def __init__(self, title, author, pages):
        self.title = title
        self.author = author
        self.pages = pages

    def summary(self):
        return f"Title: {self.title} | Author: {self.author} | Pages: {self.pages}"
```

<!-- grade_id: cell-d0743ed244cc9f5d, points: 1, grade: true, solution: false -->
```python
# Tests
b = Book("The Hobbit", "J.R.R. Tolkien", 310)
assert b.summary() == "Title: The Hobbit | Author: J.R.R. Tolkien | Pages: 310"
```

## Exercise 3: Add a Read Tracker

**Task**: Update the `Book` class so that:

1. It has an additional attribute `pages_read`, initialized to 0.
2. Add a method `read(pages)` that adds the number of pages read to `pages_read`, but **never allows it to exceed the total number of pages**.
3. Add a method `progress()` that returns the percentage of the book read as an integer.

Example:
```python
b = Book("Short Book", "Me", 100)
b.read(30)
b.progress() → 30
b.read(100)
b.progress() → 100

<details>
  <summary><strong>Click here for a hint!</strong></summary>
  Use <code>min()</code> inside <code>read()</code> to ensure <code>pages_read</code> never exceeds <code>pages</code>. In <code>progress()</code>, use integer division or <code>int()</code> with the formula: <code>(pages_read / pages) * 100</code>.
</details>

<!-- grade_id: cell-f602a7dbd4c85dec, points: 0, grade: false, solution: true -->
```python
class Book:
    def __init__(self, title, author, pages):
        self.title = title
        self.author = author
        self.pages = pages
        self.pages_read = 0

    def summary(self):
        return f"Title: {self.title} | Author: {self.author} | Pages: {self.pages}"

    def read(self, pages):
        self.pages_read = min(self.pages_read + pages, self.pages)

    def progress(self):
        return int((self.pages_read / self.pages) * 100)
```

<!-- grade_id: cell-b43f1eea77f21d78, points: 1, grade: true, solution: false -->
```python
# Tests
b = Book("Read Me", "Anon", 200)
b.read(50)
assert b.pages_read == 50
assert b.progress() == 25
b.read(300)
assert b.pages_read == 200
assert b.progress() == 100
```

## Exercise 4: Create a `Library` Class with a Collection

**Task**: Create a class `Library` that:

1. Stores a list of `Book` objects.
2. Has a method `add_book(book)` to add a book to the library.
3. Has a method `total_pages()` that returns the sum of pages of all books in the library.

You will need to reuse the `Book` class from previous exercises.

Example:
```python
lib = Library()
lib.add_book(Book("Book A", "A", 100))
lib.add_book(Book("Book B", "B", 200))
lib.total_pages() → 300

<details>
  <summary><strong>Click here for a hint!</strong></summary>
  Keep a list in the constructor, like <code>self.books = []</code>, and use <code>append()</code> to add books. Use <code>sum()</code> with a generator expression in <code>total_pages()</code>.
</details>

<!-- grade_id: cell-7b35aa681a742567, points: 0, grade: false, solution: true -->
```python
class Library:
    def __init__(self):
        self.books = []

    def add_book(self, book):
        self.books.append(book)

    def total_pages(self):
        return sum(book.pages for book in self.books)
```

<!-- grade_id: cell-54094819dd644d09, points: 1, grade: true, solution: false -->
```python
# Tests
lib = Library()
lib.add_book(Book("A", "Author A", 100))
lib.add_book(Book("B", "Author B", 200))
assert lib.total_pages() == 300
```

## Exercise 5: Prevent Negative Page Counts

**Task**: Update the `Book` class to ensure:

- If someone tries to create a `Book` with a negative number of pages, raise a `ValueError` with the message:  
  `"Pages must be a non-negative integer."`

Example:
```python
Book("Broken Book", "Nobody", -5) → raises ValueError

<details>
  <summary><strong>Click here for a hint!</strong></summary>
  In the <code>__init__</code> method, check if <code>pages</code> is less than 0 and raise an exception if so.
</details>

<!-- grade_id: cell-d7b5b7267044189b, points: 0, grade: false, solution: true -->
```python
class Book:
    def __init__(self, title, author, pages):
        if pages < 0:
            raise ValueError("Pages must be a non-negative integer.")
        self.title = title
        self.author = author
        self.pages = pages
        self.pages_read = 0

    def summary(self):
        return f"Title: {self.title} | Author: {self.author} | Pages: {self.pages}"

    def read(self, pages):
        self.pages_read = min(self.pages_read + pages, self.pages)

    def progress(self):
        return int((self.pages_read / self.pages) * 100)
```

<!-- grade_id: cell-ff3398e553f91bb7, points: 1, grade: true, solution: false -->
```python
# Tests
try:
    Book("Bad", "None", -10)
    assert False, "Expected ValueError for negative pages"
except ValueError as e:
    assert str(e) == "Pages must be a non-negative integer."
```

## Exercise 6: Add `__str__` Representation for `Book`

**Task**: Add a `__str__` method to the `Book` class that returns a user-friendly description:

```python
f"{self.title} by {self.author} — {self.pages} pages"
```

Example:
```python
Book("Dune", "Frank Herbert", 500) → str → "Dune by Frank Herbert — 500 pages"

<details>
  <summary><strong>Click here for a hint!</strong></summary>
  Define a <code>__str__</code> method and return a formatted string using the class attributes.
</details>

<!-- grade_id: cell-bbe26668ef390083, points: 0, grade: false, solution: true -->
```python
class Book:
    def __init__(self, title, author, pages):
        if pages < 0:
            raise ValueError("Pages must be a non-negative integer.")
        self.title = title
        self.author = author
        self.pages = pages
        self.pages_read = 0

    def summary(self):
        return f"Title: {self.title} | Author: {self.author} | Pages: {self.pages}"

    def read(self, pages):
        self.pages_read = min(self.pages_read + pages, self.pages)

    def progress(self):
        return int((self.pages_read / self.pages) * 100)

    def __str__(self):
        return f"{self.title} by {self.author} — {self.pages} pages"
```

<!-- grade_id: cell-ab78b6a30975957f, points: 1, grade: true, solution: false -->
```python
# Tests
b = Book("Dune", "Frank Herbert", 500)
assert str(b) == "Dune by Frank Herbert — 500 pages"
```

## Exercise 7: Filter Books by Author

**Task**: Extend the `Library` class with a method `find_books_by_author(author)` that:

- Returns a list of `Book` objects written by the given author.

Example:
```python
lib.find_books_by_author("J.K. Rowling") → [Book("HP1", "J.K. Rowling", 300), ...]

<details>
  <summary><strong>Click here for a hint!</strong></summary>
  Use a list comprehension to filter books where <code>book.author == author</code>.
</details>

<!-- grade_id: cell-22e6d779f4d349cb, points: 0, grade: false, solution: true -->
```python
class Library:
    def __init__(self):
        self.books = []

    def add_book(self, book):
        self.books.append(book)

    def total_pages(self):
        return sum(book.pages for book in self.books)

    def find_books_by_author(self, author):
        return [book for book in self.books if book.author == author]
```

<!-- grade_id: cell-217fd750e594a28c, points: 1, grade: true, solution: false -->
```python
# Tests
lib = Library()
lib.add_book(Book("HP1", "J.K. Rowling", 300))
lib.add_book(Book("HP2", "J.K. Rowling", 350))
lib.add_book(Book("LOTR", "Tolkien", 400))

books_by_rowling = lib.find_books_by_author("J.K. Rowling")
assert len(books_by_rowling) == 2
assert all(book.author == "J.K. Rowling" for book in books_by_rowling)
```

## Exercise 8: Remove Book by Title

**Task**: Add a method `remove_book(title)` to the `Library` class that:

- Removes the first book found with the given title.
- If no book with that title exists, raise a `ValueError` with the message: `"Book not found."`

<details>
  <summary><strong>Click here for a hint!</strong></summary>
  Loop through the list of books and use <code>.remove()</code> to remove the matching one.
</details>

<!-- grade_id: cell-8bceb82aff46f69f, points: 0, grade: false, solution: true -->
```python
class Library:
    def __init__(self):
        self.books = []

    def add_book(self, book):
        self.books.append(book)

    def total_pages(self):
        return sum(book.pages for book in self.books)

    def find_books_by_author(self, author):
        return [book for book in self.books if book.author == author]

    def remove_book(self, title):
        for book in self.books:
            if book.title == title:
                self.books.remove(book)
                return
        raise ValueError("Book not found.")
```

<!-- grade_id: cell-669c655220149fd6, points: 1, grade: true, solution: false -->
```python
# Tests
lib = Library()
lib.add_book(Book("A", "Auth", 10))
lib.add_book(Book("B", "Auth", 20))
lib.remove_book("A")
assert len(lib.books) == 1
try:
    lib.remove_book("C")
    assert False
except ValueError as e:
    assert str(e) == "Book not found."
```

## Exercise 9: Sort Books by Page Count

**Task**: Add a method `sort_by_pages(descending=False)` to the `Library` class that:

- Returns a new list of books sorted by number of pages.
- If `descending=True`, the list should be sorted in descending order.

<details>
  <summary><strong>Click here for a hint!</strong></summary>
  Use <code>sorted()</code> with a lambda function as the <code>key</code>.
</details>

<!-- grade_id: cell-e0ffc7317ba18618, points: 0, grade: false, solution: true -->
```python
class Library:
    def __init__(self):
        self.books = []

    def add_book(self, book):
        self.books.append(book)

    def total_pages(self):
        return sum(book.pages for book in self.books)

    def find_books_by_author(self, author):
        return [book for book in self.books if book.author == author]

    def remove_book(self, title):
        for book in self.books:
            if book.title == title:
                self.books.remove(book)
                return
        raise ValueError("Book not found.")

    def sort_by_pages(self, descending=False):
        return sorted(self.books, key=lambda book: book.pages, reverse=descending)
```

<!-- grade_id: cell-07e501e79c60c223, points: 1, grade: true, solution: false -->
```python
# Tests
lib = Library()
lib.add_book(Book("A", "Auth", 50))
lib.add_book(Book("B", "Auth", 10))
lib.add_book(Book("C", "Auth", 30))

asc = lib.sort_by_pages()
desc = lib.sort_by_pages(descending=True)

assert [b.pages for b in asc] == [10, 30, 50]
assert [b.pages for b in desc] == [50, 30, 10]
```

## Exercise 10: Count Books by Title Keyword

**Task**: Add a method `count_books_with(keyword)` to the `Library` class that:

- Returns the number of books that contain the given keyword (case-insensitive) in their title.

<details>
  <summary><strong>Click here for a hint!</strong></summary>
  Use <code>keyword.lower() in book.title.lower()</code> to check case-insensitively.
</details>

<!-- grade_id: cell-a3f4aecc50c0387a, points: 0, grade: false, solution: true -->
```python
class Library:
    def __init__(self):
        self.books = []

    def add_book(self, book):
        self.books.append(book)

    def total_pages(self):
        return sum(book.pages for book in self.books)

    def find_books_by_author(self, author):
        return [book for book in self.books if book.author == author]

    def remove_book(self, title):
        for book in self.books:
            if book.title == title:
                self.books.remove(book)
                return
        raise ValueError("Book not found.")

    def sort_by_pages(self, descending=False):
        return sorted(self.books, key=lambda book: book.pages, reverse=descending)

    def count_books_with(self, keyword):
        return sum(1 for book in self.books if keyword.lower() in book.title.lower())
```

<!-- grade_id: cell-fce422d6ce08c532, points: 1, grade: true, solution: false -->
```python
# Tests
lib = Library()
lib.add_book(Book("Python Basics", "A", 100))
lib.add_book(Book("Advanced Python", "B", 150))
lib.add_book(Book("Java Fundamentals", "C", 200))

assert lib.count_books_with("python") == 2
assert lib.count_books_with("Java") == 1
```

## Optional Exercise 1: Define Equality Between Books

**Task**: Add a method `__eq__` to the `Book` class so that:

- Two books are considered **equal** if they have the **same title and author**.
- The number of pages does **not** affect equality.

Example:
```python
Book("1984", "Orwell", 300) == Book("1984", "Orwell", 400) → True

<details>
  <summary><strong>Click here for a hint!</strong></summary>
  Implement <code>__eq__(self, other)</code> and compare <code>self.title</code> and <code>self.author</code> with <code>other.title</code> and <code>other.author</code>.
</details>

<!-- grade_id: cell-12f9846137e033ec, points: 0, grade: false, solution: true -->
```python
class Book:
    def __init__(self, title, author, pages):
        self.title = title
        self.author = author
        self.pages = pages

    def __eq__(self, other):
        return self.title == other.title and self.author == other.author
```

<!-- grade_id: cell-d5924c11c6c10856, points: 1, grade: true, solution: false -->
```python
# Tests
b1 = Book("1984", "Orwell", 300)
b2 = Book("1984", "Orwell", 400)
b3 = Book("Animal Farm", "Orwell", 300)

assert b1 == b2
assert b1 != b3
```

## Optional Exercise 2: Use Inheritance to Define `EBook`

**Task**: Create a subclass `EBook` that inherits from `Book`, and:

- Adds a new attribute: `file_size` (in megabytes).
- Overrides the `__str__` method to add this info at the end, e.g.:

```python
"Eloquent Python by Jane Smith — 200 pages (File size: 1.2MB)"

<details>
  <summary><strong>Click here for a hint!</strong></summary>
  Use <code>super().__init__()</code> to call the parent constructor. Then add the <code>file_size</code> and override <code>__str__()</code>.
</details>

<!-- grade_id: cell-575cff159afab7fa, points: 0, grade: false, solution: true -->
```python
class EBook(Book):
    def __init__(self, title, author, pages, file_size):
        super().__init__(title, author, pages)
        self.file_size = file_size

    def __str__(self):
        return f"{self.title} by {self.author} — {self.pages} pages (File size: {self.file_size}MB)"
```

<!-- grade_id: cell-0c8807b558271e7d, points: 1, grade: true, solution: false -->
```python
# Tests
ebook = EBook("Python Tricks", "Dan", 250, 1.5)
assert isinstance(ebook, Book)
assert ebook.file_size == 1.5
assert "1.5MB" in str(ebook)
```

## Optional Exercise 3: Use a Private Attribute for Rating

**Task**: Modify the `Book` class to include a **private attribute** `__rating` (between 0 and 5). Then:

1. Add a method `set_rating(r)` to set the rating. Raise a `ValueError` if `r` is not between 0 and 5.
2. Add a method `get_rating()` to return the current rating.

This teaches basic encapsulation and data protection using private variables.

<details>
  <summary><strong>Click here for a hint!</strong></summary>
  Prefix the attribute with double underscore: <code>self.__rating</code>. Use <code>set_rating()</code> and <code>get_rating()</code> to access it.
</details>

<!-- grade_id: cell-e43fd5698eb615cd, points: 0, grade: false, solution: true -->
```python
class Book:
    def __init__(self, title, author, pages):
        self.title = title
        self.author = author
        self.pages = pages
        self.__rating = None

    def set_rating(self, r):
        if r < 0 or r > 5:
            raise ValueError("Rating must be between 0 and 5.")
        self.__rating = r

    def get_rating(self):
        return self.__rating
```

<!-- grade_id: cell-cc29ed1067cf3ee9, points: 1, grade: true, solution: false -->
```python
# Tests
b = Book("Encapsulated", "Author", 100)
b.set_rating(4.5)
assert b.get_rating() == 4.5

try:
    b.set_rating(6)
    assert False
except ValueError:
    pass
```